Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **34 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joram Erbarth"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - The Delta Median problem and Randomization

## Question 1 [time estimate: 16 minutes]

### Question 1a [time estimate: 7 minutes]

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import random
def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Parameters
    ----------
    A : list
        A list of numerical values
        
    Returns
    -------
    A : list
        Sorted list
        
    """
    #we set the variable to true to control the while loop
    B = True 
    #while b is true we shuffle the list
    while B:
        random.shuffle(A)
        #we set b to false if the lists are equal this does not get changed  
        B = False
        #we go over the list and check for each element if it is bigger than the one before
        for i in range(1,len(A)):
            if A[i] < A[i-1]:
                #we set b to true if one element is not sorted
                B = True
            
    #we return the list            
    return A
    raise NotImplementedError()
    
print(worst_sort([-5,3,7,4,0,9]))
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b [time estimate: 3 minutes]
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
The best-case complexity of this algorithm is n. This is because it takes n time to check if the list is sorted and n time to shuffle the list. Both have to be done exactly once in the best case because the first shuffle produces a sorted list. 

### Question 1c [time estimate: 3 minutes]

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
There are n! possible ways to sort the array. Each of these arrays has the same probability of being produced when shuffling the list. Therefore, the average complexity is n!. Each time the algotiyhm has a chance of 1/n! to produce a sorted array.

### Question 1d [time estimate: 3 minutes]

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
It is feasible for small inputs. For very small inputs like 2 the performance is not that different from other sorting algorithms. For larger inputs, it becomes unfeasible fast. For n=5 n! it is 120 witch is unreasonable high, and for n = 10 it is 3628800.

### [Optional] Question 1e [time estimate: 5 minutes]

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
An algorithm that sorts randomly, where each possible list's probability is not the same, could perform even worse. Depending on the exact parameters, it could be hard to sort even small lists. For example, a list of 2 elements could be hard to sort if the wrongly sorted array has a 90% probability. It would take 10 steps on average instead of 2 as seen before. For inputs larger than 2 it could be not feasible to use that algorithm. 

## Question 2 [time estimate: 18 minutes]

Approximate median finder. Given a list and $\tilde{\delta}$ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the $(50-\tilde{\delta}/2)$% and $(50+\tilde{\delta}/2)$% percentiles, where $\tilde{\delta}=100\cdot \delta$. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a [time estimate: 5 minutes]
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is 
    within (50-delta/2)% and(50+delta/2)% percentiles.
    
    X% percentile of a list is defined to be the smallest number in 
    the list that is as large as at least X% numbers in the list. 
    
    Parameters
    ----------
    A : list
        A list of numerical values
    num 
        The number of interest
    delta : float
        The error, lies between 0 and 0.5

    Returns
    -------
    isApproxMed : bool
        A boolean value indicating whether num is whithin the bound
    
    """
    #we return the number if the array has only one element
    if len(A) == 1:
        
        return True
    
    #we define the upper and the lower bound depending on delta
    upper = ((50 + (100*(delta/2)))*len(A)/100)
    lower = ((50 - (100*(delta/2)))*len(A)/100)
    
    #we count the number of smaller elements in the list
    count = 0
    for i in A:
        if num > i :
            count += 1
        
    #we compare if the count is between the bounds
    if count > lower and count < upper:
        return True
    else:
        return False
    raise NotImplementedError()
    
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0, 1], 0.4, .25) == True)
assert(check_approx_med([0, 1], 0, .25) == False)
assert(check_approx_med([0], 0, .25) == True)
```

### Question 2b [time estimate: 5 minutes]
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
import random
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 
    50-delta/2% and 50+delta/2% percentiles
    
    Paramters
    ---------
    A : list
        A list of numerical values
    Delta : int
        The error, lies between 0 and 0.5
        
    Returns
    -------
    num
        The approximated median if it is found within 100//delta trials
        (this is one possible maximum number of retries. While you are 
        encouraged to play around with another limits, the submitted work 
        must use this number)
    None
        Finding failed, if nothing is found within 100//delta trials
   
   
    Note
    ----
    The 100//delta is chosen here as it is the expected number of trials for 
    the first successful finding to occur (using geometric distribution)
    
    """
    #we define the number of tests
    for i in range(int(100//delta)):
        #we choose a random number 
        num = random.choice(A)
        #we check if the number is a median and return the number if it is
        if check_approx_med(A,num,delta)==True:
            return num
    #we return none if we could not find the median
    return None
    
    raise NotImplementedError()

```

### Question 2c [time estimate: 3 minutes]

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with $\delta$ or $N$ (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
First, we have to calculate the number of elements within delta. This is (delta*100)*N/100. This is the number of acceptable outputs. The number of wrong outputs, therefore, is 1-(delta*100/N). Because every number has the same probability of being chosen, the probability for each trial is (1-((delta*100)*N/100)/N. Since we do 100//delta trials the probability of failure after all the trials is: ((1-((delta*100)*N/100))/N)*(100//delta). We can see that this does not scale with delta because the probability for each trial increases as delta decreases, but the number of trials increases, decreasing the probability. It does not scale with N as well because the number of possibilities for each trial and the number of possible medians depends on N. 

### Question 2d [time estimate: 5 minutes]
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
The check_approx_med has a runtime of n because it has to go over the array once, and the rest of the operations take constant time. The probability that the median is found in each round is (delta*100)*N/100))/N, which can be simplified to delta. We do this 100/delta times, so we expect 100/delta  * delta = 100 iterations. Each iteration has the runtime of n because it contains the other function, so the expected runtime is 100 n.  

```python

```

```python

```