Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **63 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joram Erbarth"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Computational Applications of Randomized Algorithms

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson Heaps and Heapsort, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1 [time estimate: 10 minutes]
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq

def add_to_median_heap(minh, maxh, elem):
    if len(minh) == 0:
        minh.append(-elem)
    elif len(maxh) == 0:
        maxh.append(elem)
    
    elif -elem >= minh[0]:
        heapq.heappush(minh, -elem)
    else:
        heapq.heappush(maxh, elem)
    if len(minh) +1 > len(maxh):
        heapq.heappush(maxh, -heapq.heappop(minh))
    if len(maxh) +1 > len(minh):
        heapq.heappush(minh, -heapq.heappop(maxh))
    #raise NotImplementedError()
```

## Question 2 [time estimate: 5 minutes]
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh):
    # YOUR CODE HERE
    if len(minh) == len(maxh) +1:
        return -minh[0]
    if len(minh)+1 == len(maxh):
        return maxh[0]
    if len(minh) == len(maxh):
        return (-minh[0]+maxh[0])/2

print(median([-2,-1],[3,3]))
```

Output:
```
2.5
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
1
2.0
3
4.0
5
6.0
7
8.0
9
10.0
11
12.0
13
14.0
15
16.0
17
18.0
19
20.0
21
22.0
23
24.0
25
26.0
27
28.0
29
30.0
31
32.0
33
34.0
35
36.0
37
38.0
39
40.0
41
42.0
43
44.0
45
46.0
47
48.0
49
50.0
```

## Question 4 [time estimate: 3 minutes]
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
In the worst case scenario, we first have to make three comparisons that take constant time O(1). Then we have to push an element in a heap. This takes O(log(n)) time. In the worst case, we know how to do another comparison, another push, and a pop that also takes log(n). We could describe the worst case runtime as c 3log(n) +3O(1), which has the bound O(log(n)). 

## Question 5 [time estimate: 3 minutes] 
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
The median always runs in constant time O(1). It just compares two values that are part of the input.

## Question 6 [time estimate: 4 minutes]

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
A normal median algorithm has to sort the list and then pick the median. The heap method splits the list into heaps and picks the median. The picking takes O(1) time in both cases. The runtime of the normale medina algorithm depends on the time of the sorting. If we assume a time of (nlog(n)) as we found for quicksort, this is worse than the time of (nlog(n/2)). it is log(n/2) because we have to push or pop from a heap that at most has about half of the n elements.

## [Optional] Question 7 [time estimate: 15 minutes]

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1 [time estimate: 10 minutes]

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def qselect(lst,k):
    
    if len(lst) == 1:
        return lst[0]
    
    r = len(lst)-1
    p = 0
    x = lst[r]
    i = p-1
    #sorting according to the pivot
    for j in range(p,r,1):
        if lst[j] <= x:
            i += 1
            lst[i], lst[j] = lst[j], lst[i]
    lst[i+1], lst[r] = lst[r], lst[i+1]
    if i+1 == k:
        return lst[i+1]
    elif k > i+1:
        qselect(lst[i+1:],k-i)
    
    elif k < i+1:
        qselect(lst[:i],k-i)
    #raise NotImplementedError()
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

Output:
```
<ipython-input-78-4accc171a3e6> in qselect(lst, k)
     20 
     21     elif k < i+1:
---> 22         qselect(lst[:i],k-i)
     23     #raise NotImplementedError()

<ipython-input-78-4accc171a3e6> in qselect(lst, k)
      6     r = len(lst)-1
      7     p = 0
----> 8     x = lst[r]
      9     i = p-1
     10     #sorting according to the pivot

IndexError: list index out of range
```

## Question 2 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.

```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100)) 
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))



```

Output:
```
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
None
99
```

## Question 3 [time estimate: 5 minutes]

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
T(n) = T(n-x) + O(n)

Finding the partition takes n time, and we have to partition a list that smaller than the original list.

## Question 4 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
In the best case, the partition finds the element immediately without recursion. This takes O(n) time.

## Question 5 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
In the worst case, our new list is just one smaller than the first one:
T(n) = T(n-1) + O(n)
This can be expressed as O($N^2$) because we will partition n times in the worst case.