Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joram Erbarth"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    #first, we define c, the biggest sum, as -infinity to comparison variable 
    c = float("-inf")
    #we iterate over every number in the list 
    for i in range(len(A)):
        #we set our sum to 0
        s=0
        #we iterate over every item again from the item we are currently at
        for j in range(i,len(A)):
            #we add every new number to the sum
            s += A[j]
            #if the sum is bigger than the currently biggest sum, we change the biggest sum and remember the beginning and end of that subarray
            if s > c:
                a = i
                b = j
                c = s
    #we return the tuple
    return (a,b,c)
    raise NotImplementedError()

A = [-2, -3, 4, -1, -2, 1]
print(bruteforce_max_subarray(A))
```

Output:
```
(2, 2, 4)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
It would not be necessary to re-run the entire maximum subarray. We just have to check if any subarray that contains the new number is bigger than our current maximum. We already know that our subarray is bigger than all others, which do not contain the new number.

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
The statement is false. If B's maximum subarray does not include 8, then this subarray is also a subarray of A, because A is the same as B despite 8. If the subarray is in A, and is the maximum subarray of B, than it has to be the maximum subarray of A as well. This is because all subarrays in A are also in B. So it is not possible that the maximum subarray of A is different from B, and the new subarray does not include the new number. 

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
import math

def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    # If the length of the array is 1 we return this one number, because it has to be the maximum array. 
    if len(x) == 1:
        return B[0]
    #we implement a counter, and our current maximum
    count = 0
    max_s = float("-infinity")
    # we iterate through the whole list starting with the last item
    for i in range(len(x)-1,-1,-1):
        # we add each item to the counter
        count += x[i]
        #if the counter is larger than our current max, we set our current max to the counter 
        if count > max_s:
            max_s = count
    # if our current max is larger then the max of the previous array, we return our current max and otherwise, the max of our last array
    if max_s > mx:
        return max_s
    else:
        return mx
        
        
        
        
    raise NotImplementedError()
    
print(incremental_max_subarray(B[:1], 0))
B[0]
```

Output:
```
-2
-2
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    # we set a, our current max, to minus infinity
    a= float("-infinity")
    # we iterate over each sublist from len(1) to len(A)
    for i in range(1,len(A)+1):
        # we define the sublist 
        B=A[:i]
        # we find the max from the sublist 
        st = incremental_max_subarray(B,a)
        # if the max of the sublist is larger than our current max we set our current maxt to the max of the sublsit 
        if st > a:
            a = st
    return st
        
    
print(max_subarray([-2,1,-1,2,-5]))
print(max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]))
print(max_subarray([-2]))
print(max_subarray([1,2,3,4]))
print(max_subarray([-2, -5, 6, -2, -3, -1, 5, -6]))
```

Output:
```
2
7
-2
10
6
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
This is less efficient than the divide and conquer approch. This is, becasue this algorithm still has a time efficincy of $O(n^2)$ which is bigger than $O(n log(n))$. The function still has two nested loops. The first loop is within the function and the second within the incremental_max_subarray function. It works essentially in the same way as my implementation of the bruteforce. 

```python

```